# Stanford Online Databases Course

I have asked you to complete four of the modules from [Stanford's online databases course](https://lagunita.stanford.edu/courses/DB/2014/SelfPaced/about). 

- Introduction and Relational Databases
- XML Data
- JSON Data
- SQL

With the exception of the first module, you will earn a Statement of Accomplishment after completing each module with a score of 90% or higher. (Don't worry, you can try every exercise an unlimited number of times.) The Statement of Accomplishment is a PDF file with a verification URL at the bottom. Please copy and paste this verification URL into each appropriate cell below.

```python
import requests
from bs4 import BeautifulSoup
```

Please enter your full name and assign it to the variable `name` as a string. You should enter your full name exactly as you entered it when you created your account with Stanford Online (and exactly as it appears on your Statement of Accomplishment).

```python
name = "Wilson Leong"
```

Please assign your verification URL for the "XML Data" module to the variable `url` as a string. You can run the cell immediately below to check that you have done this correctly.

<!-- grade_id: xml, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()

url = "https://verify.lagunita.stanford.edu/SOA/ff164d4f911a40279c9930dab3e2f0e0/"
```

<!-- grade_id: xml_tests, points: 30, grade: true, solution: false -->
```python
index = requests.get(url)
soup = BeautifulSoup(index.text, 'html.parser')
form_elements = soup.findAll("dd")

assert(url[:41] == "https://verify.lagunita.stanford.edu/SOA/")
assert(form_elements[0].text == name)
assert(form_elements[1].text == url[41:-1])
assert(form_elements[2].text == "XML Data")
```

Please assign your verification URL for the "JSON Data" module to the variable `url` as a string. You can run the cell immediately below to check that you have done this correctly.

<!-- grade_id: json, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()

url = "https://verify.lagunita.stanford.edu/SOA/a6fca18868794b3d84bdf04a4a192d15/"
```

<!-- grade_id: json_tests, points: 30, grade: true, solution: false -->
```python
index = requests.get(url)
soup = BeautifulSoup(index.text, 'html.parser')
form_elements = soup.findAll("dd")

assert(url[:41] == "https://verify.lagunita.stanford.edu/SOA/")
assert(form_elements[0].text == name)
assert(form_elements[1].text == url[41:-1])
assert(form_elements[2].text == "JSON Data")
```

Please assign your verification URL for the "JSON Data"(SQL?) module to the variable `url` as a string. You can run the cell immediately below to check that you have done this correctly.

<!-- grade_id: sql, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()

url = "https://verify.lagunita.stanford.edu/SOA/331469328bf3436f98646c09b9600eee/"
```

<!-- grade_id: sql_tests, points: 30, grade: true, solution: false -->
```python
index = requests.get(url)
soup = BeautifulSoup(index.text, 'html.parser')
form_elements = soup.findAll("dd")

assert(url[:41] == "https://verify.lagunita.stanford.edu/SOA/")
assert(form_elements[0].text == name)
assert(form_elements[1].text == url[41:-1])
assert(form_elements[2].text == "SQL")
```